<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 361513  -->
  <question type="stack">
    <name>
      <text>[Modular full example] Circle (Synchronous) and Phasor Diagram</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p></p>
<p>Von einem zweipoligen Turbogenerator sind folgende Daten bekannt:</p>

<div style="max-width: 40em;">
\[
\begin{align*}
S_{\mathrm{N}} &= {@SN_MVA@} \, \mathrm{MVA} (\mathrm{Y}) \\
U_{\mathrm{N}} &= {@UN_kVY_disp@} \, \mathrm{kV(Y)} \\
\operatorname{cos}(\varphi) &= {@cosphi_disp@} \, \text{(übererregt)} \\
x_{\sigma} &= {@xsigma_disp@} \, \textrm{p.u.} \\
\dfrac{I_{\mathrm{k0}}}{I_{\mathrm{N}}} &= {@INk0_by_IN_disp@} \\
\end{align*}
\]
</div>

<p>Alle Verluste sowie Sättigungserscheinungen dürfen vernachlässigt werden.</p>


<br>
<p>a) Berechnen Sie</p>
<ul>
<li>die Nennleistung \(P_{\mathrm{N}} = \) [[input:ansPN]] [[validation:ansPN]] \(\, \mathrm{MW}\), &nbsp;[[feedback:prtPN]]</li>

<li>den Nennstrom \(I_{\mathrm{N}} = \) [[input:ansIN]] [[validation:ansIN]] \(\, \mathrm{A}\), &nbsp;[[feedback:prtIN]]</li>

<li>die Synchronreaktanz \(X_{\mathrm{d}} = \) [[input:ansXd]] [[validation:ansXd]] \(\, \Omega\), &nbsp;[[feedback:prtXd]]</li>

<li>die Streureaktanz \(X_{\mathrm{\sigma}} = \) [[input:ansXsigma]] [[validation:ansXsigma]] \(\, \Omega\) &nbsp;[[feedback:prtXsigma]] &nbsp;und</li>

<li>die Hauptreaktanz \(X_{\mathrm{h}} = \) [[input:ansXh]] [[validation:ansXh]] \(\, \Omega\). &nbsp;[[feedback:prtXh]]</li></li>

</ul>

<br>
<p>b) Konstruieren Sie die Stromortskurve für Nennerregung (Maßstab \(m_I = {@mI_Abycm@} \, \mathrm{A} / \mathrm{cm}\)) und markieren Sie diese als solche.</p>
<p>Starten Sie mit der Beschriftung der Koordinatenachsen.</p>

[[jsxgraph width="600px" height="600px" input-ref-ansAxisLabels="ansAxisLabelsRef" input-ref-ansSOK="ansSOKRef" input-ref-ansStateStorage="ansStateStorageRef" input-ref-ansBP="ansBPRef"]]

var params = {
  withVectorButton: true,
  checkSOK: true,
  checkBP: true,
};

[[include src="https://raw.githubusercontent.com/jknaut/experimental-stack-jsxgraph/refs/heads/main/components/constructionBoard.js" /]]

createConstructionBoard(params);

[[/jsxgraph]]

<div hidden>
<div>[[input:ansAxisLabels]] [[validation:ansAxisLabels]]</div>
<div>[[input:ansSOK]] [[validation:ansSOK]]</div>
<div>[[input:ansStateStorage]] [[validation:ansStateStorage]]</div>
<div>[[input:ansBP]] [[validation:ansBP]]</div>
</div>

[[feedback:prtSOK]]

<br>
<p>Ermitteln Sie</p>
<ul>
<li>den Lastwinkel \(\vartheta_{\mathrm{LN}} = \) [[input:ansThetaLN]] [[validation:ansThetaLN]] (Grad), &nbsp;[[feedback:prtThetaLN]]</li>

<li>die Polradspannung \(U_{\mathrm{PN}} = \) [[input:ansUPN]] [[validation:ansUPN]] \(\, \mathrm{kV}\) &nbsp;[[feedback:prtUPN]] &nbsp;und</li>

<li>den fiktiven Erregernennstrom \(I'_{\mathrm{EN}} = \) [[input:ansIENFiktiv]] [[validation:ansIENFiktiv]] \(\, \mathrm{kA}\). &nbsp;[[feedback:prtIENFiktiv]]</li>

</ul>

<br>
<p>c) Der Generator wird bei Betrieb am Netz belastet mit \(S = {@S_MVA@} \, \mathrm{MVA}\) bei \(\cos \varphi = {@cosphiS@}\).</p>
<p>Tragen Sie diesen Betriebspunkt <u>oben</u> in die Stromortskurve ein und markieren Sie diesen als solchen.</p>
[[feedback:prtBP]]
<br>
<p>Ist dieser Betrieb dauernd zulässig?</p>
<p>[[input:ansBzYesNo]] [[validation:ansBzYesNo]] &nbsp;<p>[[feedback:prtBzYesNo]]</p></p>
<p>Geben Sie eine Begründung an:</p>
<p>[[input:ansBzReason]] [[validation:ansBzReason]]</p>
<p>[[feedback:prtBzReason]]</p>
<p>Wie groß ist der fiktive Erregerstrom in diesem Betriebspunkt?</p>
<p>\(I'_{\mathrm{E}} = \) [[input:ansIstrichE]] [[validation:ansIstrichE]] \(\, \mathrm{kA}\) &nbsp;[[feedback:prtIstrichE]]</p>

<br>
<p>d) Zeichnen Sie das Zeigerdiagramm für den Betriebspunkt in (c).</p>
[[jsxgraph width="600px" height="600px" input-ref-ansZDAxisLabels="ansAxisLabelsRef" input-ref-ansZD="ansZDRef" input-ref-ansZDStateStorage="ansStateStorageRef"]]

var params = {
  useMathJax: true,
  withPhasorDiagram: true,
  vecLabels: {#vecLabels#},
  tapePrecision: 2,
  withVectorButton: false,
  checkSOK: false,
  checkBP: false,
};

[[include src="https://raw.githubusercontent.com/jknaut/experimental-stack-jsxgraph/refs/heads/main/components/constructionBoard.js" /]]

createConstructionBoard(params);

[[/jsxgraph]]

<div hidden>[[input:ansZDAxisLabels]] [[validation:ansZDAxisLabels]]</div>
<div hidden>[[input:ansZD]] [[validation:ansZD]]</div>
<div hidden>[[input:ansZDStateStorage]] [[validation:ansZDStateStorage]]</div>

[[feedback:prtZD]]



[[if test="is(debugMode)"]]
<details class="ideal-debug"><summary>Debug-Informationen</summary>
[[ debug /]]
</details>
[[/if]]

<p>&nbsp;</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2023121100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Created by Johannes Knaut, OTH Amberg-Weiden */
/* With requirements from Prof. Dr. Heiko Zatocil, OTH Amberg-Weiden */

/* Im Einsatz immer auf false stellen, sonst sehen die Studis Debug-Informationen */
debugMode: false;

/* Funktion zur Interpretation von Koordinaten je nach Achsenbeschriftung */
interpreteCoords(coords, csystem):=block(m1: -9999, m2: -9999, c1: csystem[1], c2: csystem[2], m1: if is(c1=1) then -coords[2] elseif is(c1=3) then coords[2] elseif is(c1=2) then -coords[1] else coords[1], m2: if is(c2=1) then -coords[2] elseif is(c2=3) then coords[2] elseif is(c2=2) then -coords[1] else coords[1], mapped: [m1, m2], return(mapped));

/*************/
/*Zweipolig!!*/

r1: rand_with_step(1,5,1);

/*Scheinleistung im Nennbetrieb in VA*/
SN_MVA: [250, 220, 230, 240, 260][r1]; /* MVA */
SN_VA: 1000 * 1000 * SN_MVA; /* VA */

/*Nenneffektivwert der Spannung in V*/
UN_kVY: [138/10, 112/10, 124/10, 142/10, 156/10][r1]; /*kV(Y)*/
UN_VY: 1000 * UN_kVY; /*V(Y)*/
UN_kVY_disp: decimalplaces(UN_kVY, 1); /*z.B. 13.8*/

/*Leistungsfaktor*/
cosphi: [-8/10, -6/10, -7/10, -8/10, -9/10][r1]; /*(übererregt)*/
cosphi_disp: decimalplaces(cosphi, 1); /*z.B. -0.8*/

/*Bezogene Streureaktanz*/
xsigma: [2/10, 1/10, 12/100, 15/100, 25/100][r1]; /*p.u.*/
xsigma_disp: decimalplaces(xsigma, 1); /*z.B. 0.2*/

/*I_N: Nennwert des Primär-/Statorstroms in A*/
/*I_k0: Kurzschlusstrom in A*/
INk0_by_IN: 63/100; /*ohne Einheit*/
INk0_by_IN_disp: decimalplaces(INk0_by_IN, 2);

/*TEILAUFGABE 1*/
/*Nennleistung: P_N = S_N \cdot cos_phi*/
PN_MW: SN_MVA * cosphi; /*MW*/

/*Nennstrom: S_N = \sqrt(3) \cdot U_N \cdot I_N*/
IN_A: SN_VA / (sqrt(3) * UN_VY); /*A*/

/*Synchronreaktanz*/
Xd_Ohm: UN_VY / (sqrt(3) * INk0_by_IN * IN_A); /*Ohm*/

/*Streureaktanz*/
ZN_Ohm: UN_VY / (sqrt(3) * IN_A); /*Ohm*/
Xsigma_Ohm: 2/10 * ZN_Ohm; /*Ohm*/
Xh_Ohm: Xd_Ohm - Xsigma_Ohm; /*Ohm*/

/*TEILAUFGABE 2*/
/*Massstab*/
mI_Abycm: 2000; /*A/cm*/

/*Induktiver Strom*/
UN_by_Xdsqrt3_A: UN_VY / (sqrt(3) * Xd_Ohm);
UN_by_Xdsqrt3_cm: UN_by_Xdsqrt3_A / mI_Abycm; /*cm*/

/*Nennstrom*/
IN_cm: IN_A / mI_Abycm; /*cm*/

/*Nennstromwinkel*/
phiN: acos(cosphi);
phiN_Grad: phiN / %pi * 180; /*Grad*/

/*Zeichnerisch: Die Länge U_PN / X_d aus Stromortskurve ablesen und umrechnen in kA*/
/*Rechnerisch: Mit Kosinussatz c^2 = a^2 + b^2 - 2 \cdot a \cdot b \cdot \cos(\gamma) */
UPN_by_Xd_cm: sqrt(UN_by_Xdsqrt3_cm^2 + IN_cm^2 - 2 * UN_by_Xdsqrt3_cm * IN_cm * cos(3/2*%pi-phiN));
UPN_by_Xd_A: UPN_by_Xd_cm * mI_Abycm; /*A*/

/* Loesung fuer SOK */
expected_sok_type: "flexcircle";
taSOKCenter: [0, -float(UN_by_Xdsqrt3_cm)];
taSOKRadius: float(UPN_by_Xd_cm);
taSOK: [expected_sok_type, taSOKCenter, taSOKRadius];

/*Zeichnerisch: Den Winkel theta_LN aus Stromortskurve ablesen*/
/*Rechnerisch: Mit Sinussatz*/
thetaLN: asin(IN_cm / UPN_by_Xd_cm * sin(3/2*%pi-phiN));
thetaLN_Grad: thetaLN / %pi * 180; /*Grad*/

/*Polradspannung*/
UPN_V: UPN_by_Xd_A * Xd_Ohm; /*V*/
UPN_kV: UPN_V / 1000; /*kV*/

/*Fiktiver Erregernennstrom*/
IstrichEN_A: UPN_V / Xh_Ohm; /*A*/
IstrichEN_kA: IstrichEN_A / 1000; /*kA*/

/*TEILAUFGABE 3*/
S_MVA: 220; /*MVA*/
S_VA: 1000 * 1000 * S_MVA; /*VA*/
cosphiS: decimalplaces(cosphi_disp+0.1, 1); /* z.B. -0.7*/

/*Ständerstrom*/
I_A: S_VA / (sqrt(3) * UN_VY); /*A*/
I_cm: I_A / mI_Abycm; /*cm*/

/*Stromwinkel*/
phiS: acos(cosphiS);
phiS_Grad: phiS / %pi * 180; /*Grad*/

/*Zeicherisch IstrichE_A */
/*Zeichnerisch 1: Ablesen der Länge zwischen Betriebspunkt und Enpunkt des induktiven Stroms*/
/*Zeichnerisch 2: Abgelesene Länge / UPN_by_Xd_cm * IstrichEN_A */
/*Rechnerisch IstrichE_A */
/*Rechnerisch: Wieder mit Kosinussatz*/
Istrich_cm: sqrt(UN_by_Xdsqrt3_cm^2 + I_cm^2 - 2 * UN_by_Xdsqrt3_cm * I_cm * cos(3/2*%pi-phiS));
IstrichE_A: Istrich_cm / UPN_by_Xd_cm * IstrichEN_A; /*A*/
IstrichE_kA: IstrichE_A / 1000; /*kA*/

/*Loesungen*/
taPN: decimalplaces(PN_MW, 0);
taIN: decimalplaces(IN_A, 0);
taXd: decimalplaces(Xd_Ohm, 3);
taXsigma: decimalplaces(Xsigma_Ohm, 3);
taXh: decimalplaces(Xh_Ohm, 3);

taThetaLN: decimalplaces(thetaLN_Grad, 1);
taUPN: decimalplaces(UPN_kV, 1);
taIENFiktiv: decimalplaces(IstrichEN_kA, 1);
bpComplex: I_cm * %e^(%i*(phiS));
bpCoords: float([realpart(bpComplex), imagpart(bpComplex)]);
taBP: ["flexvector", [0, 0], bpCoords];

taStateStorageDefault: "{'objMap':[],'selectedID':null,'sokID':'-1','bpID':'-1'}";
taStateStorage: sconcat("{\"objMap\":[[\"id0084b3c4158d1\",{\"type\":\"flexcircle\",\"id\":\"id0084b3c4158d1\",\"coordsCenter\":", interpreteCoords(taSOKCenter,[1,4]), ",\"radius\":", taSOKRadius,"}],[\"idff16c04b3464f8\",{\"type\":\"flexvector\",\"id\":\"idff16c04b3464f8\",\"coordsFirst\":[0,0],\"coordsSecond\":[-3.28653346969,-3.22144715417]}]],\"selectedID\":\"idff16c04b3464f8\",\"sokID\":\"id0084b3c4158d1\",\"bpID\":\"idff16c04b3464f8\"}");

taddBzYesNo: [[1, true, "Ja"], [2, false, "Nein"]];
taBzYesNo: first(mcq_correct(taddBzYesNo));
taBzReason: "Betriebspunkt zulässig da innerhalb des erlaubten Strombereichs der Stromortskurve."
taIstrichE: decimalplaces(IstrichE_kA, 1);

/*Globale Hilfsvariablen*/
aTol: 0.5001;



/*Zeigerdiagramm*/

/*Spannungs-Maßstab*/
mU_Vbycm: float(ZN_Ohm * mI_Abycm);

I1_cx: float(I_A * %e^(%i*phiS));

U1_VY: decimalplaces(UN_VY / sqrt(3), 2);
U1cart_ZD: interpreteCoords([U1_VY/mU_Vbycm,0],[4,3]);

I1cart_ZD: interpreteCoords(bpCoords,[4,3]);

jXSigi: %i * xsigma * I1_cx;
jXSigi_ZD: interpreteCoords([realpart(jXSigi)/mU_Vbycm, imagpart(jXSigi)/mU_Vbycm], [4, 3]);

jXHi: %i * Xh_Ohm * I1_cx;
jXHi_ZD: interpreteCoords([realpart(jXHi)/mU_Vbycm, imagpart(jXHi)/mU_Vbycm], [4, 3]);

jXdi: %i * Xd_Ohm * I1_cx;
jXdi_ZD: interpreteCoords([realpart(jXdi)/mU_Vbycm, imagpart(jXdi)/mU_Vbycm], [4, 3]);

URcart_ZD: U1cart_ZD - jXSigi_ZD;

/*Polradspannung*/
UPcart_ZD: URcart_ZD - jXHi_ZD;

UR_cx: interpreteCoords(URcart_ZD,[4,3])[1] + interpreteCoords(URcart_ZD,[4,3])[2] * %i;

/*TODO: hier besser zeichnerisch versuchen, Ungenauigkeit evtl. wegen vernachlaessigtem R? */
Imu_cx: UR_cx / (%i * Xh_Ohm) * (mU_Vbycm / mI_Abycm);
Imu_ZD: interpreteCoords([realpart(Imu_cx), imagpart(Imu_cx)], [4, 3]);


vecLabels: ["\\(\\underline{U}_1\\)", "\\(\\underline{I}_1\\)", "\\(\\underline{U}_P\\)", "\\(\\underline{I}_{\\mu}\\)", "\\(\\underline{I}'_E\\)", "\\(\\underline{U}_r\\)", "\\(-j X_{1\\sigma} \\cdot \\underline{I}_1\\)", "\\(-j X_h \\cdot \\underline{I}_1\\)"];

taZD: float([[0,0],U1cart_ZD,[0,0],I1cart_ZD,[0,0],UPcart_ZD,[0,0],Imu_ZD,I1cart_ZD,Imu_ZD,[0,0],URcart_ZD,U1cart_ZD,URcart_ZD,URcart_ZD,UPcart_ZD]);

taZDMuster: float([[0,0],interpreteCoords(U1cart_ZD,[1,4]),
[0,0],interpreteCoords(I1cart_ZD,[1,4]),
[0,0],interpreteCoords(UPcart_ZD,[1,4]),
[0,0],interpreteCoords(Imu_ZD,[1,4]),
interpreteCoords(I1cart_ZD,[1,4]),interpreteCoords(Imu_ZD,[1,4]),
[0,0],interpreteCoords(URcart_ZD,[1,4]),
interpreteCoords(U1cart_ZD,[1,4]),interpreteCoords(URcart_ZD,[1,4]),
interpreteCoords(URcart_ZD,[1,4]),interpreteCoords(UPcart_ZD,[1,4])]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text></text>
    </specificfeedback>
    <questionnote>
      <text><![CDATA[Rand: {@r1@}<br>
<div style="max-width: 40em;">
\[
\begin{align*}
S_{\mathrm{N}} &= {@SN_MVA@} \, \mathrm{MVA} (\mathrm{Y}) \\
U_{\mathrm{N}} &= {@UN_kVY_disp@} \, \mathrm{kV(Y)} \\
\operatorname{cos}(\varphi) &= {@cosphi_disp@} \, \text{(übererregt)} \\
x_{\sigma} &= {@xsigma_disp@} \, \textrm{p.u.} \\
\dfrac{I_{\mathrm{k0}}}{I_{\mathrm{N}}} &= {@INk0_by_IN_disp@} \\
\end{align*}
\]
</div>
U_Maßstab: {@mU_Vbycm@}<br>
I_Maßstab: {@mI_Abycm@}<br>


[[jsxgraph width="600px" height="600px" input-ref-ansAxisLabels="ansAxisLabelsRef" input-ref-ansSOK="ansSOKRef" input-ref-ansStateStorage="ansStateStorageRef" input-ref-ansBP="ansBPRef"]]

var params = {
  withVectorButton: true,
  checkSOK: true,
  checkBP: true,
  showStaticSynchronousSOK: true,
  sok_ctr: {#interpreteCoords(taSOKCenter,[1,4])#},
  sok_rad: {#taSOKRadius#},
};

[[include src="https://raw.githubusercontent.com/jknaut/experimental-stack-jsxgraph/refs/heads/main/components/constructionBoard.js" /]]

createConstructionBoard(params);

[[/jsxgraph]]


[[jsxgraph width="600px" height="600px" input-ref-ansZDAxisLabels="ansAxisLabelsRef" input-ref-ansZD="ansZDRef" input-ref-ansZDStateStorage="ansStateStorageRef"]]

var params = {
  useMathJax: true,
  withPhasorDiagram: true,
  showPhasorSolution: true,
  vecPositions: {#taZDMuster#},
  vecLabels: {#vecLabels#},
  tapePrecision: 2,
  withVectorButton: false,
  checkSOK: false,
  checkBP: false,
};

[[include src="https://raw.githubusercontent.com/jknaut/experimental-stack-jsxgraph/refs/heads/main/components/constructionBoard.js" /]]

createConstructionBoard(params);

[[/jsxgraph]]]]></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span>&nbsp;Ihre Antwort ist richtig.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span>&nbsp;Ihre Antwort ist teilweise korrekt.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Falsche Antwort.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ansAxisLabels</name>
      <type>algebraic</type>
      <tans>[1,4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansBP</name>
      <type>algebraic</type>
      <tans>taBP</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint><![CDATA[["-1", [], []]]]></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansBzReason</name>
      <type>string</type>
      <tans>taBzReason</tans>
      <boxsize>45</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>Kurze Begründung eingeben...</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansBzYesNo</name>
      <type>dropdown</type>
      <tans>taddBzYesNo</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansIENFiktiv</name>
      <type>numerical</type>
      <tans>taIENFiktiv</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansIN</name>
      <type>numerical</type>
      <tans>taIN</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansIstrichE</name>
      <type>numerical</type>
      <tans>taIstrichE</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansPN</name>
      <type>numerical</type>
      <tans>taPN</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansSOK</name>
      <type>algebraic</type>
      <tans>taSOK</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint><![CDATA[["-1", [], -1]]]></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansStateStorage</name>
      <type>string</type>
      <tans>taStateStorage</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint><![CDATA[""]]></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansThetaLN</name>
      <type>algebraic</type>
      <tans>taThetaLN</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>°, Grad</forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansUPN</name>
      <type>numerical</type>
      <tans>taUPN</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansXd</name>
      <type>numerical</type>
      <tans>taXd</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansXh</name>
      <type>numerical</type>
      <tans>taXh</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansXsigma</name>
      <type>numerical</type>
      <tans>taXsigma</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansZD</name>
      <type>algebraic</type>
      <tans>taZDMuster</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[[-9,11],[-6,11],[-9,10],[-6,10],[-9,9],[-6,9],[-9,8],[-6,8],[-9,7],[-6,7],[-9,6],[-6,6],[-9,5],[-6,5],[-9,4],[-6,4]]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansZDAxisLabels</name>
      <type>algebraic</type>
      <tans>[1,4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ansZDStateStorage</name>
      <type>string</type>
      <tans><![CDATA[""]]></tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint><![CDATA[""]]></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prtBP</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[axisLabelsIncomplete: member(0, ansAxisLabels);
axisLabelsValid: not(axisLabelsIncomplete);

strictTol: 0.18;
relativeTol: 0.3; /* 5 Prozent Abweichung erlauben */
studentBPType: ansBP[1];
studentBPFirst: ansBP[2];
studentBPSecond: ansBP[3];

teacherBPFirst: interpreteCoords([0, 0], ansAxisLabels);
teacherBPSecond: interpreteCoords(bpCoords, ansAxisLabels);

bpWasChosen: not(studentBPType="-1");
bpValidType: is(studentBPType="flexvector");

bpCorrectFirst: is(ATNumAbsolute(studentBPFirst, teacherBPFirst, strictTol)[2]);
bpCorrectSecond: is(ATNumAbsolute(studentBPSecond, teacherBPSecond, strictTol)[2]);]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>axisLabelsValid</sans>
        <tans>true</tans>
        <testoptions>0.5001</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtBP-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBP-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Bitte beschriften Sie beide Koordinatenachsen.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>bpWasChosen</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtBP-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBP-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben noch kein geometrisches Objekt als Betriebspunkt markiert.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>bpValidType</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prtBP-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBP-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Stellen Sie den Betriebspunkt als Ortsvektor dar.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>bpCorrectFirst</sans>
        <tans>true</tans>
        <testoptions>0.3001</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prtBP-4-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBP-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Bei der Darstellung des Betriebspunktes muss der Vektor vom Ursprung ausgehen.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>bpCorrectSecond</sans>
        <tans>true</tans>
        <testoptions>0.3001</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtBP-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sie haben den Betriebspunkt richtig eingezeichnet.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBP-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben den Betriebspunkt nicht richtig eingezeichnet.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtBzReason</name>
      <value>0.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>0</feedbackstyle>
      <feedbackvariables>
        <text>mansBzReason: ansBzReason;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ansBzYesNo</sans>
        <tans>taBzYesNo</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtBzReason-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<details><summary><b>Aufdecken</b></summary>
{@taBzReason@}
</details>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBzReason-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtBzYesNo</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ansBzYesNo</sans>
        <tans>taBzYesNo</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtBzYesNo-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtBzYesNo-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtIENFiktiv</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansIENFiktiv</sans>
        <tans>taIENFiktiv</tans>
        <testoptions>aTol</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtIENFiktiv-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtIENFiktiv-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtIN</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansIN</sans>
        <tans>taIN</tans>
        <testoptions>50.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtIN-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtIN-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtIstrichE</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansIstrichE</sans>
        <tans>taIstrichE</tans>
        <testoptions>0.5001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtIstrichE-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtIstrichE-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtPN</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansPN</sans>
        <tans>taPN</tans>
        <testoptions>3.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtPN-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtPN-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtSOK</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[axisLabelsIncomplete: member(0, ansAxisLabels);
axisLabelsValid: not(axisLabelsIncomplete);

strictTol: 0.18;
relativeTol: 0.03; /* 5 Prozent Abweichung erlauben */
studentSOKType: ansSOK[1];
studentSOKCenter: ansSOK[2];
studentSOKRadius: ansSOK[3];

teacherSOKCenter: interpreteCoords(taSOKCenter, ansAxisLabels);
teacherSOKRadius: taSOKRadius;

sokWasChosen: not(studentSOKType="-1");
sokValidType: is(studentSOKType="flexcircle");

sokCorrectCenter: is(ATNumAbsolute(studentSOKCenter, teacherSOKCenter, strictTol)[2]);
sokCorrectRadius: is(ATNumRelative(studentSOKRadius, teacherSOKRadius, relativeTol)[2]);]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>axisLabelsValid</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtSOK-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtSOK-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Bitte beschriften Sie beide Koordinatenachsen.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>sokWasChosen</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtSOK-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtSOK-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sie haben noch kein Objekt als Stromortskurve markiert.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>sokValidType</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prtSOK-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Sie haben korrekt bestimmt, dass die Stromortskurve in diesem Fall ein Kreis ist.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtSOK-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Die Stromortskurve sollte in diesem Fall {@if is(studentSOKType="flexvector") then "kein Vektor" else "keine Gerade"@} sein, sondern ein anderes geometrisches Objekt.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>sokCorrectCenter</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prtSOK-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Der Mittelpunkt Ihrer Stromortskurve ist an der richtigen Stelle.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prtSOK-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Der Mittelpunkt Ihrer Stromortskurve ist an der falschen Stelle.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>sokCorrectRadius</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtSOK-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Der Radius Ihrer Stromortskurve ist ebenfalls korrekt.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtSOK-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Der Radius Ihrer Stromortskurve ist nicht korrekt.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>sokCorrectRadius</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.4</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtSOK-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Der Radius Ihrer Stromortskurve ist korrekt.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtSOK-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Der Radius Ihrer Stromortskurve ist nicht korrekt.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtThetaLN</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansThetaLN</sans>
        <tans>taThetaLN</tans>
        <testoptions>aTol</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtThetaLN-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtThetaLN-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtUPN</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansUPN</sans>
        <tans>taUPN</tans>
        <testoptions>aTol</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtUPN-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtUPN-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtXd</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansXd</sans>
        <tans>taXd</tans>
        <testoptions>0.1001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtXd-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtXd-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtXh</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansXh</sans>
        <tans>taXh</tans>
        <testoptions>0.1001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtXh-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtXh-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtXsigma</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>NumAbsolute</answertest>
        <sans>ansXsigma</sans>
        <tans>taXsigma</tans>
        <testoptions>0.1001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtXsigma-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtXsigma-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtZD</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[axisLabelsIncomplete: member(0, ansZDAxisLabels);
axisLabelsValid: not(axisLabelsIncomplete);

pair_elements(lst) := block(
    [result: []],
    for i:1 step 2 thru length(lst) do
        result: endcons([lst[i], lst[i+1]], result),
    return(result)
);

studentList: pair_elements(maplist(lambda([x], float(x)), ansZD));
teacherList: pair_elements(maplist(lambda([x], interpreteCoords(float(x), ansZDAxisLabels)), taZD));


stList: zip_with("[", studentList, teacherList);
zeroTol: 0.1;
angleTol: 0.06; /*ca. 3 Grad*/

fThree(x):=[if is(ATNumAbsolute(x[1][1], x[2][1], 0.25)[2]=true) then "&check;" else "<b>&times;</b>", if is(ATNumAbsolute(mat_norm(matrix(x[1][2]-x[1][1]), frobenius), mat_norm(matrix(x[2][2]-x[2][1]), frobenius), 0.25)[2]=true) then "&check;" else "<b>&times;</b>", if is(ATNumAbsolute(carg((x[1][2]-x[1][1])[1]+%i*(x[1][2]-x[1][1])[2]), carg((x[2][2]-x[2][1])[1]+%i*(x[2][2]-x[2][1])[2]), angleTol)[2]=true) then "&check;" else "<b>&times;</b>"];

fstList: maplist(fThree, stList);
vfstList: map(cons, vecLabels, fstList);

allCorr: every(lambda([x], is(x=["&check;","&check;","&check;"])), fstList);
allFalse: every(lambda([x], is(x=["<b>&times;</b>","<b>&times;</b>","<b>&times;</b>"])), fstList);]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>axisLabelsValid</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtZD-1-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtZD-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Bitte beschriften Sie zuerst beide Koordinatenachsen.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>allCorr</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtZD-2-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prtZD-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>allFalse</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prtZD-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prtZD-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtZD-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<style>
table#vfst {
            /*width: 50%;*/
            margin: 20px auto;
            border-collapse: collapse;
        }
table#vfst th, td {
            padding: 10px;
            border: 1px solid #ddd;
            text-align: center;
        }
table#vfst th {
            background-color: #f4f4f4;
        }
table#vfst tr:nth-child(odd) {
            background-color: #f9f9f9;
        }
table#vfst tr:nth-child(even) {
            background-color: #fff;
        }
</style>
<table id="vfst">
    <tr><th>Vektor</th><th>Startpunkt</th><th>Länge</th><th>Neigung</th></tr>
[[foreach row="args(vfstList)"]]
    <tr>[[foreach cell="row"]]<td>{@cell@}</td>[[/foreach]]</tr>
[[/foreach]]
</table>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtZD-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>8</deployedseed>
    <deployedseed>16</deployedseed>
    <deployedseed>17</deployedseed>
    <deployedseed>28</deployedseed>
    <deployedseed>29</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Testfall, der annimmt, dass die Eingabe der Trainer/innen die volle Punktezahl erreicht.</description>
      <testinput>
        <name>ansAxisLabels</name>
        <value>[1,4]</value>
      </testinput>
      <testinput>
        <name>ansBP</name>
        <value>taBP</value>
      </testinput>
      <testinput>
        <name>ansBzReason</name>
        <value>taBzReason</value>
      </testinput>
      <testinput>
        <name>ansBzYesNo</name>
        <value>first(mcq_correct(taddBzYesNo))</value>
      </testinput>
      <testinput>
        <name>ansIENFiktiv</name>
        <value>taIENFiktiv</value>
      </testinput>
      <testinput>
        <name>ansIN</name>
        <value>taIN</value>
      </testinput>
      <testinput>
        <name>ansIstrichE</name>
        <value>taIstrichE</value>
      </testinput>
      <testinput>
        <name>ansPN</name>
        <value>taPN</value>
      </testinput>
      <testinput>
        <name>ansSOK</name>
        <value>taSOK</value>
      </testinput>
      <testinput>
        <name>ansStateStorage</name>
        <value><![CDATA[""]]></value>
      </testinput>
      <testinput>
        <name>ansThetaLN</name>
        <value>taThetaLN</value>
      </testinput>
      <testinput>
        <name>ansUPN</name>
        <value>taUPN</value>
      </testinput>
      <testinput>
        <name>ansXd</name>
        <value>taXd</value>
      </testinput>
      <testinput>
        <name>ansXh</name>
        <value>taXh</value>
      </testinput>
      <testinput>
        <name>ansXsigma</name>
        <value>taXsigma</value>
      </testinput>
      <expected>
        <name>prtBP</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtBP-5-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtBzReason</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtBzReason-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtBzYesNo</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtBzYesNo-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtIENFiktiv</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtIENFiktiv-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtIN</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtIN-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtIstrichE</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtIstrichE-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtPN</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtPN-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtSOK</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtSOK-6-F</expectedanswernote>
      </expected>
      <expected>
        <name>prtThetaLN</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtThetaLN-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtUPN</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtUPN-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtXd</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtXd-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtXh</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtXh-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prtXsigma</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtXsigma-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>